# Chinese Restaurant Analysis

## Yelp Data Set
Yelp provides a dataset that includes over 150,000 businesses in many different areas and almost 7 million reviews for those businesses. This dataset can be used to answer many questions related to the success of businesses with regard to different factors. The Yelp star review as well as the number of reviews written for the business can be positive or negative indicators of success. 

```python
import pandas as pd
business = pd.read_csv("yelp_business.csv", encoding= 'unicode_escape')
business.head()
```

Output:
```
business_id                      name  \
0  Pns2l4eNsfO8kk83dixA6A  Abby Rappoport, LAC, CMQ   
1  mpf3x-BjTdTEA3yCZrAYPw             The UPS Store   
2  tUFrWirKiKi_TAnsVWINQQ                    Target   
3  MTSW4McQd7CbVtyjqoe9mw        St Honore Pastries   
4  mWMc6_wTdE0EUBKIGXDVfA  Perkiomen Valley Brewery   

                           address           city state postal_code  \
0           1616 Chapala St, Ste 2  Santa Barbara    CA       93101   
1  87 Grasso Plaza Shopping Center         Affton    MO       63123   
2             5255 E Broadway Blvd         Tucson    AZ       85711   
3                      935 Race St   Philadelphia    PA       19107   
4                    101 Walnut St     Green Lane    PA       18054   

    latitude   longitude  stars  review_count  is_open  \
0  34.426679 -119.711197    5.0             7        0   
1  38.551126  -90.335695    3.0            15        1   
2  32.223236 -110.880452    3.5            22        0   
3  39.955505  -75.155564
```

```python
import pandas as pd
review = pd.read_csv("yelp_reviews.csv", encoding= 'unicode_escape')
review.head()
```

Output:
```
review_id                 user_id             business_id  \
0  saUsX_uimxRlCVr67Z4Jig  8g_iMtfSiwikVnbP2etR0A  YjUWPpI6HXG530lwP-fb2A   
1  AqPFMleE6RsU23_auESxiA  _7bHUi9Uuf5__HHc_Q8guQ  kxX2SOes4o-D3ZQBkiMRfA   
2  Sx8TMOWLNuJBWer-0pcmoA  bcjbaE6dDog4jkNY91ncLQ  e4Vwtrqf-wpJfwesgvdgxQ   
3  JrIxlS1TzJ-iCu79ul40cQ  eUta8W_HdHMXPzLBBZhL1A  04UD14gamNjLY0IDYVhHJg   
4  6AxgBCNX_PNTOxmbRSwcKQ  r3zeYsv1XFBRA4dJpL78cw  gmjsEdUsKpj9Xxu6pdjH0g   

   stars  useful  funny  cool  \
0      3       0      0     0   
1      5       1      0     1   
2      4       1      0     1   
3      1       1      2     1   
4      5       0      2     0   

                                                text             date  
0  Family diner. Had the buffet. Eclectic assortm...   2/5/2014 20:30  
1  Wow!  Yummy, different,  delicious.   Our favo...    1/4/2015 0:01  
2  Cute interior and owner (?) gave us tour of up...  1/14/2017 20:54  
3  I am a long term frequent customer of this est...  9/23/2015 23
```

```python
merged = business.merge(review, left_on='business_id', right_on='business_id')
merged = merged[merged['categories'].str.contains("Food, Restaurants", na = False)]
asian = merged[merged['categories'].str.contains("Asian")]
```

```python
print(len(asian))
chinese = merged[merged['categories'].str.contains("Chinese")]
```

Output:
```
4466
```

```python
import datetime
asian['year'] = pd.DatetimeIndex(asian['date']).year
asian.head()

chinese['year'] = pd.DatetimeIndex(chinese['date']).year
chinese.head()
```

Output:
```
C:\Users\erika\AppData\Local\Temp\ipykernel_3960\4180538679.py:2: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame.
Try using .loc[row_indexer,col_indexer] = value instead

See the caveats in the documentation: https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  asian['year'] = pd.DatetimeIndex(asian['date']).year
C:\Users\erika\AppData\Local\Temp\ipykernel_3960\4180538679.py:5: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame.
Try using .loc[row_indexer,col_indexer] = value instead

See the caveats in the documentation: https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  chinese['year'] = pd.DatetimeIndex(chinese['date']).year
        business_id            name           address          city state  \
3892128      #NAME?  Saigon Cuisine  4000 Chestnut St  Philadelphia    PA   
3892129      #NAME?  Sai
```

```python
chinese[["year", "stars_y"]].groupby("year").mean()
#print(grouped)
#grouped.plot(y='stars_y')
```

Output:
```
stars_y
year          
2005  4.000000
2006  3.833333
2007  4.090909
2008  3.758621
2009  3.542857
2010  3.544118
2011  3.669565
2012  3.876623
2013  3.735931
2014  3.674286
2015  3.836735
2016  4.041284
2017  3.964286
2018  4.212996
```

This is the code generated for some of our prototype analysis.

Tests for q1 (points: [1, 1, 1, 1, 2, 2, 1, 1, 1, 1])
```
>>> isinstance(q1_1_1, str)
True
```
```
>>> import numpy
>>> isinstance(q1_1_2, int) or isinstance(q1_1_2, numpy.int64)
True
```
```
>>> from auto_checks import check_q1_1_1
>>> check_q1_1_1(q1_1_1)
'PASS'
```
```
>>> from auto_checks import check_q1_1_2
>>> check_q1_1_2(q1_1_2)
'PASS'
```
```
>>> import numpy
>>> isinstance(q1_2, int) or isinstance(q1_2, numpy.int64)
True
```
```
>>> from auto_checks import check_q1_2
>>> check_q1_2(q1_2)
'PASS'
```
```
>>> import pandas
>>> isinstance(q1_3, pandas.core.series.Series)
True
```
```
>>> q1_3.shape
(5,)
```
```
>>> from auto_checks import check_q1_3_province
>>> check_q1_3_province(q1_3)
'PASS'
```
```
>>> from auto_checks import check_q1_3_death
>>> check_q1_3_death(q1_3)
'PASS'
```


Tests for q2 (points: [1, 1, 1, 1, 1, 1])
```
>>> import pandas
>>> isinstance(q2_1, pandas.core.series.Series)
True
```
```
>>> q2_1.shape
(5,)
```
```
>>> from auto_checks import check_q2_1_state
>>> check_q2_1_state(q2_1)
'PASS'
```
```
>>> from auto_checks import check_q2_1_date
>>> check_q2_1_date(q2_1)
'PASS'
```
```
>>> import pandas
>>> isinstance(q2_2, pandas.core.series.Series)
True
```
```
>>> q2_2.shape
(364,)
```


Tests for q3 (points: [1, 1, 1, 1, 1, 1, 1, 1])
```
>>> import pandas
>>> isinstance(q3_1, pandas.core.frame.DataFrame)
True
```
```
>>> q3_1.shape
(5, 2)
```
```
>>> print(q3_1.index.name, sorted(q3_1.columns))
rating ['F', 'M']
```
```
>>> from auto_checks import check_q3_1
>>> check_q3_1(q3_1)
'PASS'
```
```
>>> import pandas
>>> isinstance(q3_2, pandas.core.frame.DataFrame)
True
```
```
>>> q3_2.shape
(7, 2)
```
```
>>> print(q3_2.index.name, sorted(q3_2.columns))
decade ['F', 'M']
```
```
>>> from auto_checks import check_q3_2
>>> check_q3_2(q3_2)
'PASS'
```


Tests for q4 (points: [1, 1, 1, 2, 1, 1, 1, 1, 1, 1, 1, 1])
```
>>> q4_0
100000
```
```
>>> import numpy
>>> isinstance(q4_1, float) or isinstance(q4_1, numpy.float64)
True
```
```
>>> 0 <= q4_1 <= 1
True
```
```
>>> from auto_checks import check_q4_1
>>> check_q4_1(q4_1)
'PASS'
```
```
>>> import pandas
>>> isinstance(q4_2, pandas.core.series.Series)
True
```
```
>>> q4_2.shape
(5,)
```
```
>>> from auto_checks import check_q4_2_index
>>> check_q4_2_index(q4_2)
'PASS'
```
```
>>> from auto_checks import check_q4_2_rating
>>> check_q4_2_rating(q4_2)
'PASS'
```
```
>>> import pandas
>>> isinstance(q4_3, pandas.core.series.Series)
True
```
```
>>> q4_3.shape
(5,)
```
```
>>> from auto_checks import check_q4_3_index
>>> check_q4_3_index(q4_3)
'PASS'
```
```
>>> from auto_checks import check_q4_3_rating
>>> check_q4_3_rating(q4_3)
'PASS'
```
