Star Rating Analysis


```python
# Run but do not modify this code
import pandas as pd
business = pd.read_csv("yelp_business.csv", encoding= 'unicode_escape')
business.head()
```

```python
# Run but do not modify this code

import seaborn as sns
sns.set_theme()
review = pd.read_csv("yelp_reviews.csv", encoding= 'unicode_escape')
review.head()
```

```python
merged = business.merge(review, left_on='business_id', right_on='business_id')
merged = merged[merged['categories'].str.contains("Food, Restaurants", na = False)]
merged['year'] = pd.DatetimeIndex(merged['date']).year
asian = merged[merged['categories'].str.contains("Asian")]
```

Indian, Thai, Italian, Mexican, Chinese

```python
indian = merged[merged['categories'].str.contains("Indian")]
thai = merged[merged['categories'].str.contains("Thai")]
italian = merged[merged['categories'].str.contains("Italian")]
mexican = merged[merged['categories'].str.contains("Mexican")]
chinese = merged[merged['categories'].str.contains("Chinese")]
```

```python
import matplotlib.pyplot as plt

starRank = merged['stars_y'].value_counts().sort_index(ascending=False)
starRank = (starRank/(len(merged)))*100
print(starRank)



indianStar = indian['stars_y'].value_counts().sort_index(ascending=False).apply(lambda x : x / len(indian)*100)
thaiStar = thai['stars_y'].value_counts().sort_index(ascending=False).apply(lambda x : x / len(thai)*100)
italianStar = italian['stars_y'].value_counts().sort_index(ascending=False).apply(lambda x : x / len(italian)*100)
mexicanStar = mexican['stars_y'].value_counts().sort_index(ascending=False).apply(lambda x : x / len(mexican)*100)
chineseStar = chinese['stars_y'].value_counts().sort_index(ascending=False).apply(lambda x : x / len(chinese)*100)
print(indianStar)

star_graph = pd.concat([indianStar.rename('Indian'), thaiStar.rename('Thai'), italianStar.rename('Italian'),
          mexicanStar.rename('Mexican'), chineseStar.rename('Chinese')],axis=1).plot.bar(rot=0, figsize=(20,10))

bruh = pd.concat([indianStar.rename('Indian'), thaiStar.rename('Thai'), italianStar.rename('Italian'),
          mexicanStar.rename('Mexican'), chineseStar.rename('Chinese')])
print(bruh)
                 
                 
star_graph.set_xlabel("Stars")
star_graph.set_ylabel("Percentage")
```

```python
isOpen = merged[merged['is_open'] == 1]
isClosed = merged[merged['is_open'] == 0]
openStars = isOpen['stars_y'].value_counts()
openStarsRank = (openStars/(len(isOpen)))*100
closedStars = isClosed['stars_y'].value_counts()
closedStarsRank = (closedStars/(len(isClosed)))*100

import matplotlib.pyplot as plt

print(openStarsRank)
print(len(isOpen))
print(closedStarsRank)
print(len(isClosed))
print(merged[['is_open']])
op = merged[['is_open']]

total = len(op)
is_one = len(op[op['is_open'] == 1])
is_zero = len(op[op['is_open'] == 0])

data = [is_one/total, is_zero/total]

labels = ['is open', 'is not open']

print((is_one/total)*100)
print((is_zero/total)*100)


plt.pie(data, labels = labels, colors = ('lavender','cornsilk'), 
        wedgeprops={"linewidth": 1, "edgecolor": "black", 'antialiased': True})
plt.show()


#thing = merged.groupby(['is_open']).plot(kind='pie', subplots = True, figsize=(5, 5))

```

```python
import numpy as np
from scipy import stats

### chinese ci 
cStar = chinese['stars_y']
mu = np.mean(cStar)
sigma = np.std(cStar)
n = chinese.shape[0]

c_ci = stats.norm.interval(alpha=.95, loc=mu, scale = sigma/np.sqrt(n))
print("Chinese restaurant mean confidence interval:", c_ci)

### thai ci 
tStar = thai['stars_y']
mu2 = np.mean(tStar)
sigma2 = np.std(tStar)
n2 = thai.shape[0]

t_ci = stats.norm.interval(alpha=.95, loc=mu2, scale = sigma2/np.sqrt(n2))
print("Thai restaurant mean confidence interval:", t_ci)

### italian ci 
itStar = italian['stars_y']
mu3 = np.mean(cStar)
sigma3 = np.std(cStar)
n3 = italian.shape[0]

it_ci = stats.norm.interval(alpha=.95, loc=mu3, scale = sigma3/np.sqrt(n3))
print("Italian restaurant mean confidence interval:", it_ci)

### indian ci
inStar = indian['stars_y']
mu4 = np.mean(cStar)
sigma4 = np.std(cStar)
n4 = indian.shape[0]

in_ci = stats.norm.interval(alpha=.95, loc=mu4, scale = sigma4/np.sqrt(n4))
print("Indian restaurant mean confidence interval:", in_ci)

### mexican ci 
mStar = mexican['stars_y']
mu5 = np.mean(cStar)
sigma5 = np.std(cStar)
n5 = mexican.shape[0]

m_ci = stats.norm.interval(alpha=.95, loc=mu5, scale = sigma5/np.sqrt(n5))
print("Mexican restaurant mean confidence interval:", m_ci)
```

Tests for q1 (points: [1, 1, 1, 1, 2, 2, 1, 1, 1, 1])
```
>>> isinstance(q1_1_1, str)
True
```
```
>>> import numpy
>>> isinstance(q1_1_2, int) or isinstance(q1_1_2, numpy.int64)
True
```
```
>>> from auto_checks import check_q1_1_1
>>> check_q1_1_1(q1_1_1)
'PASS'
```
```
>>> from auto_checks import check_q1_1_2
>>> check_q1_1_2(q1_1_2)
'PASS'
```
```
>>> import numpy
>>> isinstance(q1_2, int) or isinstance(q1_2, numpy.int64)
True
```
```
>>> from auto_checks import check_q1_2
>>> check_q1_2(q1_2)
'PASS'
```
```
>>> import pandas
>>> isinstance(q1_3, pandas.core.series.Series)
True
```
```
>>> q1_3.shape
(5,)
```
```
>>> from auto_checks import check_q1_3_province
>>> check_q1_3_province(q1_3)
'PASS'
```
```
>>> from auto_checks import check_q1_3_death
>>> check_q1_3_death(q1_3)
'PASS'
```


Tests for q2 (points: [1, 1, 1, 1, 1, 1])
```
>>> import pandas
>>> isinstance(q2_1, pandas.core.series.Series)
True
```
```
>>> q2_1.shape
(5,)
```
```
>>> from auto_checks import check_q2_1_state
>>> check_q2_1_state(q2_1)
'PASS'
```
```
>>> from auto_checks import check_q2_1_date
>>> check_q2_1_date(q2_1)
'PASS'
```
```
>>> import pandas
>>> isinstance(q2_2, pandas.core.series.Series)
True
```
```
>>> q2_2.shape
(364,)
```


Tests for q3 (points: [1, 1, 1, 1, 1, 1, 1, 1])
```
>>> import pandas
>>> isinstance(q3_1, pandas.core.frame.DataFrame)
True
```
```
>>> q3_1.shape
(5, 2)
```
```
>>> print(q3_1.index.name, sorted(q3_1.columns))
rating ['F', 'M']
```
```
>>> from auto_checks import check_q3_1
>>> check_q3_1(q3_1)
'PASS'
```
```
>>> import pandas
>>> isinstance(q3_2, pandas.core.frame.DataFrame)
True
```
```
>>> q3_2.shape
(7, 2)
```
```
>>> print(q3_2.index.name, sorted(q3_2.columns))
decade ['F', 'M']
```
```
>>> from auto_checks import check_q3_2
>>> check_q3_2(q3_2)
'PASS'
```


Tests for q4 (points: [1, 1, 1, 2, 1, 1, 1, 1, 1, 1, 1, 1])
```
>>> q4_0
100000
```
```
>>> import numpy
>>> isinstance(q4_1, float) or isinstance(q4_1, numpy.float64)
True
```
```
>>> 0 <= q4_1 <= 1
True
```
```
>>> from auto_checks import check_q4_1
>>> check_q4_1(q4_1)
'PASS'
```
```
>>> import pandas
>>> isinstance(q4_2, pandas.core.series.Series)
True
```
```
>>> q4_2.shape
(5,)
```
```
>>> from auto_checks import check_q4_2_index
>>> check_q4_2_index(q4_2)
'PASS'
```
```
>>> from auto_checks import check_q4_2_rating
>>> check_q4_2_rating(q4_2)
'PASS'
```
```
>>> import pandas
>>> isinstance(q4_3, pandas.core.series.Series)
True
```
```
>>> q4_3.shape
(5,)
```
```
>>> from auto_checks import check_q4_3_index
>>> check_q4_3_index(q4_3)
'PASS'
```
```
>>> from auto_checks import check_q4_3_rating
>>> check_q4_3_rating(q4_3)
'PASS'
```
